```python
# Initialize Otter
import otter
grader = otter.Notebook("lab0.ipynb")
```

# Lab 0 - Intro to Labs

Welcome. The chief intent of these exercises is to apply and think about the week's material in a different way. Any technical skills acquired are incidental; practical value is not guaranteed.  

Do start with the [basic tutorial](../../tutorial/Basics.ipynb) if you are unfamiliar with the structure of Jupyter notebooks. 

## Objectives

By the end of this lab, students should be able to 
  - represent vectors as NumPy `array`s
  - perform vector operations in Python
  - cite some use cases for convex combinations

```python
# Import needed libraries at the outset. 

import numpy as np
import matplotlib.pyplot as plt
import ipywidgets as widgets

from numpy import sin, cos, pi, sqrt, exp, log, array, linspace, arange
import matplotlib.image as mpimg
from matplotlib.animation import FuncAnimation
from matplotlib.colors import hsv_to_rgb

from IPython.display import HTML, Markdown

# Uncomment this if you are using Dark Mode
plt.style.use('default')
# plt.style.use('dark_background')

rng_seed = 90210
```

## Vectors

In the basics tutorial, we have only looked at scalars and scalar operations. To move up to vectors, we will invoke the [NumPy](https://www.numpy.org/)  (pronounce however you feel is most fun) library. We invoked it above, using the conventional abbreviation `np`. We can then access its vast wealth of mathematical structures using `np.xxxxx` notation.

For vectors, we are going to use a structure called a "NumPy array" which is actually a bit more generic and can be used for things like matrices and tensors as well, but we will stick to vectors. 

A vector or $\vec v = \langle 1,2,4\rangle$ is declared this way.

```python
v = np.array([1,2,4])
```

**Note.** Note the the double sets of delimeters. We pass a list (tuple) object to the array function to make a vector. 

### Why not just use lists?

The expression `[1,2,4]` is a valid Python object (a list), but there is a good reason we don't use those as vectors. Observe:

```python
[1,2,4] + [7,6,8]
```

Useful for sure, but maybe not what we want mathematically. Python knows to reat numpy arrays as numerical objects

```python
w = np.array([7,6,8])
```

Now we can compute with `v` and `w`. 

```python
v+w
```

### Vector operations

We can accomplish our basic vector operations with relative ease. 

_We haven't seen all of these yet, so don't panic, but we will shortly._

|mathematics     | Python  | decripion     |
|----------------|----------------|----------------|
| $\vec v + \vec w $ | `v + w`| vector addition |
| $\vec v - \vec w $ | `v - w`| vector subtraction |
| $ - \vec w $ | `-w`|  negation | 
| $c \vec v $ | `c * w`| scalar multiplication |
| $\vec v \cdot \vec w $ | `np.dot(v,w)`| dot product |
| $\vec v \times \vec w $ | `np.cross(v,w)`| cross product |

```python
v + w
```

```python
v - w
```

```python
-w
```

```python
6 * v
```

```python
# it is smart enough
v * 3
```

#### Question 1

1. Find the position vector for the midpoint of positions $\vec u$ and $\vec w$ and assign it to the variable `mid_point`. 
2. Find the position vector for the point 90% of the way from position $\vec w$ to position $\vec v$ and assign it to `ninety_pct`.

**WARNING**. Use the variables, not the numeric values, in your expressions as the random seed will change in the grader. 

```python
np.random.seed(rng_seed)

v, w = np.random.rand(2, 3)

mid_point = ...
ninety_pct = ...
```

```python
grader.check("q1")
```

#### Question 2

Find the centroid of the triangle with vertices at positions $\vec u$, $\vec v$, and $\vec w$. Store this as the variable `uvw_centroid`. (*The **centroid** of a triangle is the intersection of its medians, which it can be shown always is $2/3$ the way from a vertex to the midpoint of the opposite side.*) 

```python
np.random.seed(rng_seed)

u, v, w = np.random.rand(3, 3)

uvw_centroid= ...

```

```python
grader.check("q2")
```

## Plotting vectors

One of the most popular plotting packpages for Python is [matplotlib](https://matplotlib.org/). It is powerful, but pretty far from being intuitive to use. We will wade in slowly. 

```python
## This is an import block for matplotlib. Normally we will do this in the first cell of the notebook.

%matplotlib inline
import matplotlib.pyplot as plt
```

```python
# Draw a vector from the origin to (1,2)

plt.arrow(0,0,1,2);
```

```python
# Draw a vector from the origin to (1,2), in blue, make the window bigger, and label axes.

plt.arrow(0,0,1,2,color='b',length_includes_head=True,head_width=.1);
plt.xlim([-3,3])
plt.ylim([-3,3])
plt.xlabel("$x$")
plt.ylabel("$y$")
plt.title("A Vector");
```

That's a bit better, but a lot of code for a simple thing. Let's make our lives a bit easier and define a function to do this.

```python
def plot_vector(v,base=(0,0),**kwargs):
    """Plots a vector `v` with tail at the point `base` (defaults to origin)."""
    return plt.arrow(base[0],base[1],v[0],v[1],length_includes_head=True,head_width=.2,**kwargs)
```

```python
v = np.array([-1.5,2.5])
w = np.array([3,4])

plt.gca().set_aspect('equal', adjustable='box')

plt.xlim([-5,5])
plt.ylim([-5,5])
plt.xlabel("$x$")
plt.ylabel("$y$")
plt.title("A Vector");
plt.grid(True)

plot_vector(v,color='r')
plot_vector(v,base=(3,-1),color='b')
a3 = plot_vector(v,base=(-2,2),color='purple')
```

```python
plt.gca().set_aspect('equal', adjustable='box')
plot_vector(v,color='r')
plot_vector(w,color='b')
plot_vector(v+w,color='purple')
plot_vector(v - w, base=w); # semicolon on last line suppresses printing return value.
```

## Convex Combinations

A **convex combination** of vectors is a linear combination in which all the scaling factors add to 1. For two vectors, that means something of the form $$(1 - t) \vec v + t \vec w.$$

We could think of this as a *weighted average*, a *linear interpolation* or, alternatively as a parameterized path from $\vec v$ to $\vec w$ as $t$ goes from 0 to 1.

### Linear segments

If $\vec v$ and $\vec w$ are positions of endpoints, we can descibe the positions of the line segment between them with the combination above. 

The code below illustrates this. 

```python
# plot the line segment
plt.plot([v[0], w[0]], [v[1], w[1]])

for t in linspace(0, 1, 15):
    plot_vector((1 - t)*v + t*w, color='purple')

plot_vector(v,color='r')
plot_vector(w,color='b')

plt.gca().set_aspect('equal')
```

So convex combinations are good for literally describing straight paths (line segments) in space, but more abstractly, you can think of them as a linear transition from one state to another. 

### Images

To some extent, digital images are just themselves big vectors. We could thus use convex combinations to turn a frog into a prince. 

```python
frog = mpimg.imread("frog.png")
prince = mpimg.imread("prince.png")
```

```python
@widgets.interact
def _(t = (0,1,.01)):
    plt.imshow((1 - t)*frog + t*prince)
```

*Note: The files here were prepped to compatible in in terms of dimension and format. It's not quite this simple for any two images.*

### Animation

We call a convex combination of the form above a **parameterization**. We'll hear this over and over. The parameter $t$ can represent all sorts of things. One is time. 

Let's make a simple animation of a triangle moving from one position to another. 

The key piece here is the `update` function, which draws the new position of the function based on a parameter. 

```python
fig, ax = plt.subplots(figsize=(5,5))

ax.set_xlim((0, 10))
ax.set_ylim((0, 10))

# Initial position
u, v, w = array(((0,0), (2,0), (1,2)))

# Final position
u1, v1, w1 = array(((7,5), (9,5), (8,7)))

# Set a list of blank polygons
polys = plt.fill([], alpha = .5)

# Pick the first
poly = polys[0]

def update(t):
    """Draw the polygon based on parameter value t."""
    poly.set_xy([(1 - t)*u + t*u1,
                 (1 - t)*v + t*v1,
                 (1 - t)*w + t*w1])

def ax_init():
    """Reset axis."""
    update(0)
    
anim = FuncAnimation(fig, update, frames=linspace(0,1,72), interval=1000/24)

plt.close()

HTML(anim.to_html5_video())
```

Compelling action, I know. Make it more exciting by making a new animation. This one, we'll save as a GIF.

<!-- BEGIN QUESTION -->

#### Question 3

Change the `update` function below to make an interesting GIF of your own design. 

![Closing hexagon example](hexample.gif) 

*Don't just imitate this example, and don't worry about making it complicated. Anything original that uses linear combinations of vectors will do.*

You can add more polygons by setting `N` to a positive integer, and then within `update` edit the `n`th one's vertices by using the `polys[n].set_xy` method as in the example above. `poly[n].set_color` can optionally be used to change the color. 


```python
fig, ax = plt.subplots(figsize=(5,5))

ax.set_xlim((-2,2)) # Change the view window if you wish
ax.set_ylim((-2,2)) # Change the view window if you wish

# Turn off axis markers [OPTIONAL].
ax.set_xticks([])
ax.set_yticks([])

# Set a list of blank polygons

N = 6

polys = plt.fill(*[[]]*2*N, alpha = .5)

# Pick the first
poly = polys[0]

def update(t):
    """Draw the polygon(s) based on parameter value t."""

    poly.set_xy([
        ...
        ])
    
def ax_init():
    """Reset axis."""
    update(0)

# You can adjust the length of the animation by altering the frames argument. 
anim = FuncAnimation(fig, update, frames=linspace(0,1,72), interval=1000/24)
# anim = FuncAnimation(fig, update, frames=np.concatenate((linspace(0,4,288),linspace(4,0,288))), interval=1000/24)

plt.close()

anim.save("more_interesting.gif")

HTML("""<img src="./more_interesting.gif">""")
```

<!-- END QUESTION -->

### Maps

Many applications of mathematics study maps—not atlas kind of maps—but functions that associate each point of the plane to another point, like shifting grains of sand on a board. The dimensions make these somewhat difficult to visualize, but one way to do it is draw a figure as a set of input points and use convex combinations to show it transitioning to its output position. This gives a sense of the map's action. In other words, we graph for each input/output pair a position $$(1 - t) \mbox{ input} + t \mbox{ output}$$
and then vary the parameter $t$ from $0$ to $1$. 

**Note**. The interim values may have no physical meaning or significance other that showing where a point is going.

For example consider the following map: $$f: (x, y) \mapsto (x^2 - y^2, 2xy)$$

Let's color a grid and then animate where points go under this map.

```python
%matplotlib widget

fig, ax = plt.subplots(figsize = (7,7))

ax.set_xlim([-2,2])
ax.set_ylim([-2, 2])
ax.grid(True)

# plt.axis('equal')

t = linspace(-1,1,100)

lines = []

for i in linspace(-1, 1, 9):
    p = ax.plot(t, i * np.ones_like(t), color='b')[0]
    lines.append(p)

for i in linspace(-1, 1, 9):
    p = ax.plot(i * np.ones_like(t), t, color='r')[0]
    lines.append(p)

lines_data = [l.get_xydata() for l in lines]
# define map above
def f(v):
    x, y = v # separate components
    
    #########################
    #########################
    # Modify this line to visualize a different map.
    #########################
    #########################
    out = (x**2 - y**2, 2*x*y) 
    #########################
    #########################
    
    return np.array(out)

@widgets.interact
def _(t = (0, 1, .01)):
    for line, data in zip(lines, lines_data):
        xy = data
        line.set_data(np.column_stack([(1 - t)*v + t*f(v) for v in xy]))
        
```

Determine a function $g:\mathbb{R}^2 \to \mathbb{R}^2$ that "folds" the square $\{(x,y) : -1 \leq x \leq 1, -1 \leq y \leq 1 \}$ to the unit square as shown. 

![Animation of square folded onto itself twice](q4.gif)

```python
def g(v):
    x,y = v
    out = ...
    return np.array(out)
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

#### Feedback/Reflection

Please jot down a few sentences about your experience with this lab. Helpful information would include:
  - Did these expand or enhance your understanding of linear combinations in any way?
  - Did you have difficulty with the length or amount of technical jargon/code? 
  - Is there anything mathematically you are still confused about?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Export

Save this notebook (use the File menu above). Then execute the following cell and upload the zip file to Gradescope. You'll see the autograded score and be able to resubmit until the deadline. 

```python
import zipfile, os

ASSIGNMENT_NAME="lab0"
ZIPFILE=f"{ASSIGNMENT_NAME}_submission.zip"
CWD=os.path.curdir

zip = zipfile.ZipFile(os.path.join(CWD, ZIPFILE), mode='w')

for filename in ["more_interesting.gif", f"{ASSIGNMENT_NAME}.ipynb"]:
    zip.write(os.path.join(CWD, filename))
    
zip.close()

display(HTML(f"""Download the <a href="{os.path.join(CWD,ZIPFILE)}" download="{ZIPFILE}" target="_blank">zip file</a> and upload it to the assignment on Gradescope."""))
```

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 2)
```
>>> 
>>> # make sure your answers are numpy arrays
>>> 
>>> mid_point.shape == (3,)
True
```
failure_message: Your answer is not in the correct format.
success_message: Your answer is in the correct format. Its accuracy is determined when uploaded.
```
>>> ninety_pct.shape == (3,)
True
```
failure_message: Your answer is not in the correct format.
success_message: Your answer is in the correct format. Its accuracy is determined when uploaded.


Tests for q2 (points: 2)
```
>>> 
>>> # make sure your answers are numpy arrays
>>> 
>>> uvw_centroid.shape == (3,)
True
```
failure_message: Your answer is not in the correct format.
success_message: Your answer is in the correct format. Its accuracy is determined when uploaded.


Tests for q4 (points: 2)
```
>>> v = np.array((0,1))
>>> np.allclose(v, g(v))
True
```
```
>>> v = np.array((1,0))
>>> np.allclose(v, g(v))
True
```
```
>>> w = np.array((1,1))
>>> np.allclose(g(w), (0,0))
True
```
```
>>> w = np.array((0,0))
>>> np.allclose(g(w), (1, 1))
True
```
